1. ! [X0] : add(z,X0) = X0 [input]
2. ! [X0,X1] : add(s(X0),X1) = s(add(X0,X1)) [input]
3. ! [X0,X1] : add(X0,X1) = add(X1,X0) [input]
4. ! [X0,X1] : (p(X0,X1) <=> (s(add(X1,X1)) = X0 | add(X1,X1) = X0)) [input]
5. (! [X0] : (? [X1] : p(X0,X1) => ? [X1] : p(s(X0),X1)) & ? [X1] : p(z,X1)) => ! [X0] : ? [X1] : p(X0,X1) [input]
6. ! [X0] : ? [X1] : p(X0,X1) [input]
7. ~! [X0] : ? [X1] : p(X0,X1) [negated conjecture 6]
8. (! [X0] : (? [X1] : p(X0,X1) => ? [X2] : p(s(X0),X2)) & ? [X3] : p(z,X3)) => ! [X4] : ? [X5] : p(X4,X5) [rectify 5]
9. ! [X4] : ? [X5] : p(X4,X5) | (? [X0] : (! [X2] : ~p(s(X0),X2) & ? [X1] : p(X0,X1)) | ! [X3] : ~p(z,X3)) [ennf transformation 8]
10. ! [X4] : ? [X5] : p(X4,X5) | ? [X0] : (! [X2] : ~p(s(X0),X2) & ? [X1] : p(X0,X1)) | ! [X3] : ~p(z,X3) [flattening 9]
11. ? [X0] : ! [X1] : ~p(X0,X1) [ennf transformation 7]
12. ! [X0,X1] : ((p(X0,X1) | (s(add(X1,X1)) != X0 & add(X1,X1) != X0)) & ((s(add(X1,X1)) = X0 | add(X1,X1) = X0) | ~p(X0,X1))) [nnf transformation 4]
13. ! [X0,X1] : ((p(X0,X1) | (s(add(X1,X1)) != X0 & add(X1,X1) != X0)) & (s(add(X1,X1)) = X0 | add(X1,X1) = X0 | ~p(X0,X1))) [flattening 12]
14. ! [X0] : ? [X1] : p(X0,X1) | ? [X2] : (! [X3] : ~p(s(X2),X3) & ? [X4] : p(X2,X4)) | ! [X5] : ~p(z,X5) [rectify 10]
15. ! [X0] : (? [X1] : p(X0,X1) => p(X0,sK0(X0))) [choice axiom]
16. ? [X2] : (! [X3] : ~p(s(X2),X3) & ? [X4] : p(X2,X4)) => (! [X3] : ~p(s(sK1),X3) & ? [X4] : p(sK1,X4)) [choice axiom]
17. ? [X4] : p(sK1,X4) => p(sK1,sK2) [choice axiom]
18. ! [X0] : p(X0,sK0(X0)) | (! [X3] : ~p(s(sK1),X3) & p(sK1,sK2)) | ! [X5] : ~p(z,X5) [skolemisation 14,17,16,15]
19. ? [X0] : ! [X1] : ~p(X0,X1) => ! [X1] : ~p(sK3,X1) [choice axiom]
20. ! [X1] : ~p(sK3,X1) [skolemisation 11,19]
21. add(z,X0) = X0 [cnf transformation 1]
22. add(s(X0),X1) = s(add(X0,X1)) [cnf transformation 2]
23. add(X0,X1) = add(X1,X0) [cnf transformation 3]
24. ~p(X0,X1) | add(X1,X1) = X0 | s(add(X1,X1)) = X0 [cnf transformation 13]
25. p(X0,X1) | add(X1,X1) != X0 [cnf transformation 13]
26. p(X0,X1) | s(add(X1,X1)) != X0 [cnf transformation 13]
27. ~p(z,X5) | p(sK1,sK2) | p(X0,sK0(X0)) [cnf transformation 18]
28. ~p(s(sK1),X3) | p(X0,sK0(X0)) | ~p(z,X5) [cnf transformation 18]
29. ~p(sK3,X1) [cnf transformation 20]
30. p(s(add(X1,X1)),X1) [equality resolution 26]
31. p(add(X1,X1),X1) [equality resolution 25]
32. p(z,z) [superposition 31,21]
38. add(X3,s(X2)) = s(add(X2,X3)) [superposition 22,23]
41. p(s(add(X1,s(X1))),s(X1)) [superposition 31,22]
45. p(s(s(add(X1,X1))),s(X1)) [forward demodulation 41,38]
54. p(X0,sK0(X0)) | p(sK1,sK2) [resolution 27,32]
55. p(sK1,sK2) [resolution 54,29]
64. sK1 = s(add(sK2,sK2)) | sK1 = add(sK2,sK2) [resolution 24,55]
158. p(s(sK1),s(sK2)) | sK1 = add(sK2,sK2) [superposition 45,64]
221. ~p(z,X1) | p(X0,sK0(X0)) | sK1 = add(sK2,sK2) [resolution 158,28]
454. p(X0,sK0(X0)) | sK1 = add(sK2,sK2) [resolution 221,32]
515. sK1 = add(sK2,sK2) [resolution 454,29]
584. p(s(sK1),sK2) [superposition 30,515]
669. ~p(z,X1) | p(X0,sK0(X0)) [resolution 584,28]
882. p(X0,sK0(X0)) [resolution 669,32]
1056. $false [resolution 882,29]